```python
# Initialize Otter
import otter
grader = otter.Notebook("group08_assignment2.ipynb")
```

# Assignment 2: Pandas
## Cleaning and Exploring Data with Pandas

## Assignment Modalities -- IMPORTANT -- READ CAREFULLY

This assignment contains a number of question that are worth 1--2 points each, as seen at the table below. At each question, you need to fill in the missing code, which is denoted by `...`. In some cases, you could write a single line of code, in other cases you might want to use more lines to achieve the desired outcome. **NOTE** that you should not change any code segment except the part denoted as `...`, as it might interfere with the auto grading, see below.

The assignment will be automatically graded by running tests on the provided code. Some of the tests are available after each cell with missing code. To run the automatic grader, you need to have otter-grader installed. This is not mandatory, but advised, so that you can check the validity of your answers. **NOTE** that for the final grading we will run additional tests not available in this notebook.

The installation instructions for otter-grader are here:  
<https://otter-grader.readthedocs.io/en/latest/#installation>

**IMPORTANT**: First rename this notebook to `groupXX_assignment2.ipynb`, where `XX` should be your group number for the 2nd assignment, e.g., `01`. For the autograder to work locally, make sure that in the initialization of otter the right notebook name is passed:
```
grader = otter.Notebook("groupXX_assignment2.ipynb")
```

**Note**: Installing the otter-grader is not necessary. All tests visible to you are included in the notebook as assertions.

After finishing the assignment, you should submit this notebook named `groupXX_assignment2.ipynb`.


## Score Breakdown 
Question | Points
--- | ---
1a | 1
1b | 2
1c | 1
2a | 2
2b | 1
2ci | 1
2cii | 1
2d | 2
2e | 2
3a | 1
3bi | 2
3bii | 2
3ci | 1
3cii | 1
3ciii | 1
3civ | 1
4a | 2
4b | 3
4c | 2
Total | 29


## Assignment Objective

In this homework, we will investigate restaurant food safety scores for restaurants in San Francisco. The scores and violation information have been [made available by the San Francisco Department of Public Health](https://data.sfgov.org/Health-and-Social-Services/Restaurant-Scores-LIVES-Standard/pyih-qa8i). The main goal for this assignment is to walk through the process of Data Cleaning and EDA. 


As we clean and explore these data, you will gain practice with:
* Reading simple csv files and using Pandas
* Working with data at different levels of granularity
* Identifying the type of data collected, missing values, anomalies, etc.



## Before You Start

For each question in the assignment, please write down your answer in the answer cell(s) right below the question. 

We understand that it is helpful to have extra cells breaking down the process towards reaching your final answer. If you happen to create new cells below your answer to run codes, **NEVER** add cells between a question cell and the answer cell below it. It will cause errors when we run the autograder, and it will sometimes cause a failure to generate the PDF file.

**Important note: The local autograder tests will not be comprehensive. You can pass the automated tests in your notebook but still fail tests in the autograder.** Please be sure to check your results carefully.

Finally, unless we state otherwise, try to avoid using python for loops or list comprehensions.  The majority of this assignment can be done using builtin commands in Pandas and numpy.  


```python
import numpy as np
import pandas as pd

from pathlib import Path
import os # Used to interact with the file system
```

## Obtaining the Data

### File Systems and I/O



In general, we will focus on using python commands to investigate files.  However, it can sometimes be easier to use shell commands in your local operating system.  The following cells demonstrate how to do this.

**NOTE** Do not change the location of the data.

```python
from pathlib import Path

data_dir = Path('./data')


```

## Loading Food Safety Data

We have data, but we don't have any specific questions about the data yet. Let's focus on understanding the structure of the data; this involves answering questions such as:

* Is the data in a standard format or encoding?
* Is the data organized in records?
* What are the fields in each record?

Let's start by looking at the contents of the `data` dictionary.

## Looking At the Data Files

The following codeblocks are setup. Simply run the cells; **do not modify them**. Question 1a is where you will start to write code.


You should see five CSV files. Open up `legend.csv` to see its contents.  You should see something that looks like:

    "Minimum_Score","Maximum_Score","Description"
    0,70,"Poor"
    71,85,"Needs Improvement"
    86,90,"Adequate"
    91,100,"Good"

## Reading in and Verifying Data

Based on the above information, let's attempt to load `bus.csv`, `ins2vio.csv`, `ins.csv`, and `vio.csv` into pandas dataframes with the following names: `bus`, `ins2vio`, `ins`, and `vio` respectively.

*Note:* Because of character encoding issues one of the files (`bus`) will require an additional argument `encoding='ISO-8859-1'` when calling `pd.read_csv`. At some point in your future, you should read all about [character encodings](https://diveintopython3.problemsolving.io/strings.html).

```python
# path to directory containing data

bus = pd.read_csv(data_dir/'bus.csv', encoding='ISO-8859-1')
ins2vio = pd.read_csv(data_dir/'ins2vio.csv')
ins = pd.read_csv(data_dir/'ins.csv')
vio = pd.read_csv(data_dir/'vio.csv')

#This code is essential for the autograder to function properly. Do not edit
ins_test = ins
```

Now that you've read in the files, let's try some `pd.DataFrame` methods ([docs](https://pandas.pydata.org/pandas-docs/version/0.21/generated/pandas.DataFrame.html)).
Use the `DataFrame.head` method to show the top few lines of the `bus`, `ins`, and `vio` dataframes. To show multiple return outputs in one single cell, you can use `display()`. Currently, running the cell below will display the first few lines of the `bus` dataframe. 

```python
bus.head()
```

Output:
```
business id column                         name                 address  \
0                1000        HEUNG YUEN RESTAURANT            3279 22nd St   
1              100010        ILLY CAFFE SF_PIER 39        PIER 39  K-106-B   
2              100017  AMICI'S EAST COAST PIZZERIA             475 06th St   
3              100026               LOCAL CATERING        1566 CARROLL AVE   
4              100030             OUI OUI! MACARON  2200 JERROLD AVE STE C   

            city state postal_code     latitude    longitude  phone_number  
0  San Francisco    CA       94110    37.755282  -122.420493         -9999  
1  San Francisco    CA       94133 -9999.000000 -9999.000000   14154827284  
2  San Francisco    CA       94103 -9999.000000 -9999.000000   14155279839  
3  San Francisco    CA       94124 -9999.000000 -9999.000000   14155860315  
4  San Francisco    CA       94124 -9999.000000 -9999.000000   14159702675
```

The `DataFrame.describe` method can also be handy for computing summaries of numeric columns of our dataframes. Try it out with each of our 4 dataframes. Below, we have used the method to give a summary of the `bus` dataframe. 

```python
bus.describe()
```

Output:
```
business id column     latitude    longitude  phone_number
count         6253.000000  6253.000000  6253.000000  6.253000e+03
mean         60448.948984 -5575.337966 -5645.817699  4.701819e+09
std          36480.132445  4983.390142  4903.993683  6.667508e+09
min             19.000000 -9999.000000 -9999.000000 -9.999000e+03
25%          18399.000000 -9999.000000 -9999.000000 -9.999000e+03
50%          75685.000000 -9999.000000 -9999.000000 -9.999000e+03
75%          90886.000000    37.776494  -122.421553  1.415533e+10
max         102705.000000    37.824494     0.000000  1.415988e+10
```

Now, we perform some sanity checks for you to verify that the data was loaded with the correct structure. Run the following cells to load some basic utilities (you do not need to change these at all):

First, we check the basic structure of the data frames you created:

```python
assert all(bus.columns == ['business id column', 'name', 'address', 'city', 'state', 'postal_code',
                           'latitude', 'longitude', 'phone_number'])
assert 6250 <= len(bus) <= 6260

assert all(ins.columns == ['iid', 'date', 'score', 'type'])
assert 26660 <= len(ins) <= 26670

assert all(vio.columns == ['description', 'risk_category', 'vid'])
assert 60 <= len(vio) <= 65

assert all(ins2vio.columns == ['iid', 'vid'])
assert 40210 <= len(ins2vio) <= 40220
```

Next we'll check that the statistics match what we expect. The following are hard-coded statistical summaries of the correct data.

```python
bus_summary = pd.DataFrame(**{'columns': ['business id column', 'latitude', 'longitude'],
 'data': {'business id column': {'50%': 75685.0, 'max': 102705.0, 'min': 19.0},
  'latitude': {'50%': -9999.0, 'max': 37.824494, 'min': -9999.0},
  'longitude': {'50%': -9999.0,
   'max': 0.0,
   'min': -9999.0}},
 'index': ['min', '50%', 'max']})

ins_summary = pd.DataFrame(**{'columns': ['score'],
 'data': {'score': {'50%': 76.0, 'max': 100.0, 'min': -1.0}},
 'index': ['min', '50%', 'max']})

vio_summary = pd.DataFrame(**{'columns': ['vid'],
 'data': {'vid': {'50%': 103135.0, 'max': 103177.0, 'min': 103102.0}},
 'index': ['min', '50%', 'max']})

from IPython.display import display

print('What we expect from your Businesses dataframe:')
display(bus_summary)
print('What we expect from your Inspections dataframe:')
display(ins_summary)
print('What we expect from your Violations dataframe:')
display(vio_summary)
```

Output:
```
What we expect from your Businesses dataframe:
     business id column     latitude  longitude
min                19.0 -9999.000000    -9999.0
50%             75685.0 -9999.000000    -9999.0
max            102705.0    37.824494        0.0
What we expect from your Inspections dataframe:
     score
min   -1.0
50%   76.0
max  100.0
What we expect from your Violations dataframe:
          vid
min  103102.0
50%  103135.0
max  103177.0
```

The code below defines a testing function that we'll use to verify that your data has the same statistics as what we expect. Run these cells to define the function. The `df_allclose` function has this name because we are verifying that all of the statistics for your dataframe are close to the expected values. Why not `df_allequal`? It's a bad idea in almost all cases to compare two floating point values like 37.780435, as rounding error can cause spurious failures.

```python
"""Run this cell to load this utility comparison function that we will use in various
tests below (both tests you can see and those we run internally for grading).

Do not modify the function in any way.
"""


def df_allclose(actual, desired, columns=None, rtol=5e-2):
    """Compare selected columns of two dataframes on a few summary statistics.
    
    Compute the min, median and max of the two dataframes on the given columns, and compare
    that they match numerically to the given relative tolerance.
    
    If they don't match, an AssertionError is raised (by `numpy.testing`).
    """    
    # summary statistics to compare on
    stats = ['min', '50%', 'max']
    
    # For the desired values, we can provide a full DF with the same structure as
    # the actual data, or pre-computed summary statistics.
    # We assume a pre-computed summary was provided if columns is None. In that case, 
    # `desired` *must* have the same structure as the actual's summary
    if columns is None:
        des = desired
        columns = desired.columns
    else:
        des = desired[columns].describe().loc[stats]

    # Extract summary stats from actual DF
    act = actual[columns].describe().loc[stats]

    return np.allclose(act, des, rtol)
```

We will now explore each file in turn, including determining its granularity and primary keys and exploring many of the variables individually. Let's begin with the businesses file, which has been read into the `bus` dataframe.

<br/><br/><br/>

---

<br/><br/><br/>

## Question 1a: Examining the Business Data File

From its name alone, we expect the `bus.csv` file to contain information about the restaurants. Let's investigate the granularity of this dataset.

```python
bus.head()
```

Output:
```
business id column                         name                 address  \
0                1000        HEUNG YUEN RESTAURANT            3279 22nd St   
1              100010        ILLY CAFFE SF_PIER 39        PIER 39  K-106-B   
2              100017  AMICI'S EAST COAST PIZZERIA             475 06th St   
3              100026               LOCAL CATERING        1566 CARROLL AVE   
4              100030             OUI OUI! MACARON  2200 JERROLD AVE STE C   

            city state postal_code     latitude    longitude  phone_number  
0  San Francisco    CA       94110    37.755282  -122.420493         -9999  
1  San Francisco    CA       94133 -9999.000000 -9999.000000   14154827284  
2  San Francisco    CA       94103 -9999.000000 -9999.000000   14155279839  
3  San Francisco    CA       94124 -9999.000000 -9999.000000   14155860315  
4  San Francisco    CA       94124 -9999.000000 -9999.000000   14159702675
```

The `bus` dataframe contains a column called `business id column` which probably corresponds to a unique business id.  However, we will first rename that column to `bid` for simplicity.

**Note**: In practice we might want to do this renaming when the table is loaded but for grading purposes we will do it here.


```python
bus = bus.rename(columns={"business id column": "bid"})
```

Examining the entries in `bus`, is the `bid` unique for each record (i.e. each row of data)? Your code should compute the answer, i.e. don't just hard code `True` or `False`.

Hint: use `value_counts()` or `unique()` to determine if the `bid` series has any duplicates.

```python
is_bid_unique = bus['bid'].is_unique

is_bid_unique
```

Output:
```
True
```

```python
assert(is_bid_unique)
```

```python
grader.check("q1a")
```

Output:
```
q1a results: All test cases passed!
```

<br/><br/><br/>

---


## Question 1b

We will now work with some important fields in `bus`.

1. Assign `top_names` to a list containing the top 5 most frequently used business names, from most frequent to least frequent.
2. Assign `top_addresses` to a list containing the top 5 addressses where businesses are located, from most popular to least popular.

Hint: you may find `value_counts()` helpful.

```python
top_names = bus['name'].value_counts().head(5).index.tolist()
top_addresses = bus['address'].value_counts().head(5).index.tolist()

top_names, top_addresses
```

Output:
```
(["Peet's Coffee & Tea",
  'Starbucks Coffee',
  "McDonald's",
  'Jamba Juice',
  'STARBUCKS'],
 ['Off The Grid', '428 11th St', '2948 Folsom St', '3251 20th Ave', 'Pier 41'])
```

```python
assert len(top_names) == 5
assert top_names[0] == "Peet's Coffee & Tea"
```

```python
grader.check("q1b")
```

Output:
```
q1b results: All test cases passed!
```

<br/><br/><br/>

---


## Question 1c

Based on the above exploration, what does each record represent?

A. "One location of a restaurant."
B. "A chain of restaurants."
C. "A city block."

Answer in the following cell. Your answer should be a string, either `"A"`, `"B"`, or `"C"`.

```python
# What does each record represent?  Valid answers are:
#    A "One location of a restaurant."
#    B "A chain of restaurants."
#    C "A city block."
q1c = ...
```

```python
assert(q1c in set(["A", "B", "C"]))
```

```python
grader.check("q1c")
```

<br/><br/><br/>

---

<br/><br/><br/>

# 2: Cleaning the Business Data Postal Codes

The business data contains postal code information that we can use to aggregate the ratings over regions of the city.  Let's examine and clean the postal code field.

*NOTE*: In USA, postal codes have 5 digits, like `90210`. There is also a more specific or full postal code that has four additional digits, like `90210-1007`.

The postal code (sometimes also called a ZIP code) partitions the city into regions:

<img src="https://www.usmapguide.com/wp-content/uploads/2019/03/printable-san-francisco-zip-code-map.jpg" alt="ZIP Code Map" style="width: 600px">

<br/><br/><br/>

---


## Question 2a

How many restaurants are in each ZIP code? 

In the cell below, create a **series** where the index is the postal code and the value is the number of records with that postal code in descending order of count. You may need to use `groupby()`, `size()`, or `value_counts()`. Do you notice any odd/invalid zip codes?

<!--
BEGIN QUESTION
name: q2a
points: 2
-->

<!-- grade_id: cell-d2151d673e6c36a1, points: 0, grade: false, solution: true -->
```python
zip_counts = bus['postal_code'].value_counts()
print(zip_counts)
```

Output:
```
postal_code
94103         562
94110         555
94102         456
94107         408
94133         398
             ... 
94120           1
92672           1
95105           1
941033148       1
94123-3106      1
Name: count, Length: 63, dtype: int64
```

```python
assert(type(zip_counts) == pd.Series)
assert(zip_counts["94103"] == 562)
```

```python
grader.check("q2a")
```

Output:
```
q2a results: All test cases passed!
```

<br/><br/><br/>

--- 

## Question 2b

Answer the following question about the `postal_code` column in the `bus` dataframe.

   
1. What Python data type is used to represent a ZIP code?  
    A. `str`  
    B. `int`  
    C. `bool`  
    D. `float`

*Note*: ZIP codes and postal codes are the same thing.

Please write your answers in the cell below. Your answer should be a string, either `"A"`, `"B"`, `"C"`, or `"D"`.

```python
# Check the actual type of postal_code entries
postal_code_type = type(bus['postal_code'].iloc[0])

# Map each type to its corresponding letter
if postal_code_type == str:
    q2b = "A"
elif postal_code_type == int:
    q2b = "B"
elif postal_code_type == bool:
    q2b = "C"
elif postal_code_type == float:
    q2b = "D"
```

```python
assert(q2b in set(["A", "B", "C"]))
```

```python
grader.check("q2b")
```

Output:
```
q2b results: All test cases passed!
```

<br/><br/><br/>

--- 

## Question 2c

In question 2a we noticed a large number of potentially invalid ZIP codes (e.g., "Ca").  These are likely due to data entry errors.  To get a better understanding of the potential errors in the zip codes we will:

1. Import a list of valid San Francisco ZIP codes by using `pd.read_json` to load the file `data/sf_zipcodes.json` and extract a **series** of type `str` containing the valid ZIP codes.  *Hint: set `dtype` when invoking `read_json`.*
1. Construct a `DataFrame` containing only the businesses which DO NOT have valid ZIP codes.  You will probably want to use the `Series.isin` function. 


**Step 1**

<!--
BEGIN QUESTION
name: q2ci
points: 1
-->

```python
valid_zips = pd.read_json(data_dir/'sf_zipcodes.json', typ='series', dtype=str)

valid_zips.head()
```

Output:
```
zip_codes    ['94102', '94103', '94104', '94105', '94107', ...
dtype: object
```

```python
assert( (valid_zips.dtype == object) | (valid_zips.dtype == pd.StringDtype())) 
assert(type(valid_zips) == pd.Series)
```

```python
grader.check("q2ci")
```

Output:
```
q2ci results: All test cases passed!
```

**Step 2**

<!--
BEGIN QUESTION
name: q2cii
points: 1
-->

```python
# has_valid_zip should be a boolean array
# A True value would indicate the business has a valid ZIP code

has_valid_zip = bus['postal_code'].isin(valid_zips)
invalid_zip_bus = bus[~has_valid_zip]
invalid_zip_bus.head(20)
```

Output:
```
bid                                     name  \
0     1000                    HEUNG YUEN RESTAURANT   
1   100010                    ILLY CAFFE SF_PIER 39   
2   100017              AMICI'S EAST COAST PIZZERIA   
3   100026                           LOCAL CATERING   
4   100030                         OUI OUI! MACARON   
5   100036                          Hula Truck (#2)   
6   100039                 GENKI CREPES & MINI MART   
7   100041                           UNCLE LEE CAFE   
8   100055                            Twirl and Dip   
9   100058                              SF PITA HUB   
10  100059                           DUMPLING ALLEY   
11  100069                             Mission Blue   
12  100072                  SUBWAY SANDWICHES #7307   
13  100079                           POSITIVE FOODS   
14  100081     THE MATTERHORN RESTAURANT AND BAKERY   
15  100082                                SLN CTRNG   
16  100083                     THE EPICUREAN TRADER   
17  100084       
```

```python
assert(type(invalid_zip_bus) == pd.DataFrame) 
```

```python
grader.check("q2cii")
```

Output:
```
q2cii results: All test cases passed!
```

<br/><br/><br/>

--- 

## Question 2d

In the previous question, many of the businesses had a common invalid postal code that was likely used to encode a MISSING postal code.  Do they all share a potentially "interesting address"?

In the following cell, construct a **series** that counts the number of businesses at each `address` that have this single likely MISSING postal code value.  Order the series in descending order by count. 

After examining the output, please answer the following question (2e) by filling in the appropriate variable. If we were to drop businesses with MISSING postal code values would a particular class of business be affected? 

```python
missing_zip_address_count = ...
missing_zip_address_count
```

```python
assert(type(missing_zip_address_count) == pd.Series)
assert(missing_zip_address_count['3914 Judah St'] == 1)
```

```python
grader.check("q2d")
```

<!-- END QUESTION -->

<br/><br/><br/>

--- 

## Question 2e

Examine the `invalid_zip_bus` dataframe we computed above and look at the businesses that DO NOT have the special MISSING ZIP code value. Some of the invalid postal codes are just the full 9 digit code rather than the first 5 digits. Create a new column named `postal5` in the original `bus` dataframe which contains only the first 5 digits of the `postal_code` column.

Then, for any of the `postal5` ZIP code entries that were not a valid San Fransisco ZIP Code (according to `valid_zips`), the provided code will set the `postal5` value to `None`.  

**Do not modify the provided code!**

```python
bus['postal5'] = ...


bus.loc[~bus['postal5'].isin(valid_zips), 'postal5'] = None
# Checking the corrected postal5 column
bus.loc[invalid_zip_bus.index, ['bid', 'name', 'postal_code', 'postal5']]
```

```python
assert('postal5' in bus.columns) 
assert((bus['postal5'].str.len() != 5).sum() == 221)
```

```python
grader.check("q2e")
```

<br/><br/><br/>

---

<br/><br/><br/>

# 3: Investigate the Inspection Data

Let's now turn to the inspection DataFrame. Earlier, we found that `ins` has 4 columns named 
`iid`, `score`, `date` and `type`.  In this section, we determine the granularity of `ins` and investigate the kinds of information provided for the inspections. 

Let's start by looking again at the first 5 rows of `ins` to see what we're working with.

```python
ins.head(5)
```

Output:
```
iid                    date  score                   type
0  100010_20190329  03/29/2019 12:00:00 AM     -1       New Construction
1  100010_20190403  04/03/2019 12:00:00 AM    100  Routine - Unscheduled
2  100017_20190417  04/17/2019 12:00:00 AM     -1          New Ownership
3  100017_20190816  08/16/2019 12:00:00 AM     91  Routine - Unscheduled
4  100017_20190826  08/26/2019 12:00:00 AM     -1  Reinspection/Followup
```

<br/><br/><br/>

---

## Question 3a

The column `iid` probably corresponds to an inspection id.  Is it a primary key?  Write an expression (line of code) that evaluates to `True` or `False` based on whether all the values are unique.

**Hint:** This is a very similar question to Question 1a.

```python
is_ins_iid_a_primary_key = ins['iid'].is_unique
print(is_ins_iid_a_primary_key)
```

Output:
```
True
```

```python
assert(type(is_ins_iid_a_primary_key) == bool or type(is_ins_iid_a_primary_key) == np.bool) 
```

```python
grader.check("q3a")
```

Output:
```
q3a results: All test cases passed!
```

<br/><br/><br/>

---

## Question 3b

The column `iid` appears to be the composition of two numbers and the first number looks like a business id.  

**Part 1.**: Create a new column called `bid` in the `ins` dataframe containing just the business id.  You will want to use `ins['iid'].str` operations to do this.  Also be sure to convert the type of this column to `int`

**Part 2.**: Then compute how many values in this new column are invalid business ids (i.e. do not appear in the `bus['bid']` column). This is verifying a foreign key relationship. Consider using the `pd.Series.isin` function.


**No python `for` loops or list comprehensions required!**

**Part 1**

```python
ins['bid'] = ins['iid'].str.split('_').str[0].astype(int)
ins.head()
```

Output:
```
iid                    date  score                   type  \
0  100010_20190329  03/29/2019 12:00:00 AM     -1       New Construction   
1  100010_20190403  04/03/2019 12:00:00 AM    100  Routine - Unscheduled   
2  100017_20190417  04/17/2019 12:00:00 AM     -1          New Ownership   
3  100017_20190816  08/16/2019 12:00:00 AM     91  Routine - Unscheduled   
4  100017_20190826  08/26/2019 12:00:00 AM     -1  Reinspection/Followup   

      bid  
0  100010  
1  100010  
2  100017  
3  100017  
4  100017
```

```python
assert('bid' in ins.columns) 
assert(ins['bid'].dtype == int)
```

```python
grader.check("q3bi")
```

Output:
```
q3bi results: All test cases passed!
```

**Part 2**

```python
invalid_bid_count = (~ins['bid'].isin(bus['bid'])).sum()

print(invalid_bid_count)
```

Output:
```
0
```

```python
assert(invalid_bid_count == 0) 
```

```python
grader.check("q3bii")
```

Output:
```
q3bii results: All test cases passed!
```

<br/><br/><br/>

---

## Question 3c

What if we are interested in a time component of the inspection data?  We need to examine the date column of each inspection. 

**Part 1:** What is the type of the individual `ins['date']` entries? You may want to grab the very first entry and use the `type` function in python. 

**Part 2:** Use `pd.to_datetime` to create a new `ins['timestamp']` column containing of `pd.Timestamp` objects.  These will allow us to do more date manipulation.

**Part 3:** What are the earliest and latest dates in our inspection data?  *Hint: you can use `min` and `max` on dates of the correct type.*

**Part 4:** We probably want to examine the inspections by year. Create an additional `ins['year']` column containing just the year of the inspection.  Consider using `pd.Series.dt.year` to do this.  

**Note: If you get a `SettingWithCopyWarning`, consider using the [`assign` method](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.assign.html) to add a column to a DataFrame.**

**No python `for` loops or list comprehensions required!**

**Part 1**

```python
ins_date_type = type(ins['date'][0])
ins_date_type
```

Output:
```
str
```

```python
assert(type(ins_date_type) == type) 
assert(ins_date_type in [type("str"), type("int")])
```

```python
grader.check("q3ci")
```

Output:
```
q3ci results: All test cases passed!
```

**Part 2**

```python
ins['timestamp'] = pd.to_datetime(ins['date'])
ins.head()
```

Output:
```
/var/folders/85/r49m8fjj4nb3hxp8b6pflngc0000gn/T/ipykernel_48853/4022488954.py:1: UserWarning: Could not infer format, so each element will be parsed individually, falling back to `dateutil`. To ensure parsing is consistent and as-expected, please specify a format.
  ins['timestamp'] = pd.to_datetime(ins['date'])
               iid                    date  score                   type  \
0  100010_20190329  03/29/2019 12:00:00 AM     -1       New Construction   
1  100010_20190403  04/03/2019 12:00:00 AM    100  Routine - Unscheduled   
2  100017_20190417  04/17/2019 12:00:00 AM     -1          New Ownership   
3  100017_20190816  08/16/2019 12:00:00 AM     91  Routine - Unscheduled   
4  100017_20190826  08/26/2019 12:00:00 AM     -1  Reinspection/Followup   

      bid  timestamp  
0  100010 2019-03-29  
1  100010 2019-04-03  
2  100017 2019-04-17  
3  100017 2019-08-16  
4  100017 2019-08-26
```

```python
assert(type(ins['timestamp'][1]) == pd.Timestamp) 
```

```python
grader.check("q3cii")
```

Output:
```
q3cii results: All test cases passed!
```

**Part 3**

```python
earliest_date = ins['timestamp'].min()
latest_date = ins['timestamp'].max()

print("Earliest Date:", earliest_date)
print("Latest Date:", latest_date)
```

Output:
```
Earliest Date: 2016-10-04 00:00:00
Latest Date: 2019-11-28 00:00:00
```

```python
assert(type(earliest_date) == pd.Timestamp) 
assert(type(latest_date) == pd.Timestamp)
```

```python
grader.check("q3ciii")
```

Output:
```
q3ciii results: All test cases passed!
```

**Part 4**

```python
ins['year'] = ins['timestamp'].dt.year

ins.head()
```

Output:
```
iid                    date  score                   type  \
0  100010_20190329  03/29/2019 12:00:00 AM     -1       New Construction   
1  100010_20190403  04/03/2019 12:00:00 AM    100  Routine - Unscheduled   
2  100017_20190417  04/17/2019 12:00:00 AM     -1          New Ownership   
3  100017_20190816  08/16/2019 12:00:00 AM     91  Routine - Unscheduled   
4  100017_20190826  08/26/2019 12:00:00 AM     -1  Reinspection/Followup   

      bid  timestamp  year  
0  100010 2019-03-29  2019  
1  100010 2019-04-03  2019  
2  100017 2019-04-17  2019  
3  100017 2019-08-16  2019  
4  100017 2019-08-26  2019
```

```python
assert('year' in ins.columns) 
```

```python
grader.check("q3civ")
```

Output:
```
q3civ results: All test cases passed!
```

<br/><br/><br/>

---

<br/><br/><br/>

# 4: Joining Data Across Tables

In this question we will start to connect data across mulitple tables.  We will be using the `merge` function. 

<br/><br/><br/>

--- 

## Question 4a

Let's figure out which restaurants had the lowest scores. Before we proceed, let's filter out missing scores from `ins` so that negative scores don't influence our results.

```python
ins = ins[ins["score"] > 0]
```

We'll start by creating a new dataframe called `ins_named`. It should be exactly the same as `ins`, except that it should have the name and address of every business, as determined by the `bus` dataframe. If a `business_id` in `ins` does not exist in `bus`, the name and address should be given as `NaN`. 

*Hint*: Use the merge method to join the `ins` dataframe with the appropriate portion of the `bus` dataframe. See the official [documentation](https://pandas.pydata.org/pandas-docs/stable/user_guide/merging.html) on how to use `merge`.

*Note*: For quick reference, a pandas 'left' join keeps the keys from the left frame, so if `ins` is the left frame, all the keys from `ins` are kept and if a set of these keys don't have matches in the other frame, the columns from the other frame for these "unmatched" key rows contains NaNs.

```python
ins_named = ins.merge(bus[['bid', 'name', 'address']], on='bid', how='left')

ins_named.head()
```

Output:
```
iid                    date  score                   type  \
0  100010_20190403  04/03/2019 12:00:00 AM    100  Routine - Unscheduled   
1  100017_20190816  08/16/2019 12:00:00 AM     91  Routine - Unscheduled   
2  100041_20190520  05/20/2019 12:00:00 AM     83  Routine - Unscheduled   
3  100055_20190425  04/25/2019 12:00:00 AM     98  Routine - Unscheduled   
4  100055_20190912  09/12/2019 12:00:00 AM     82  Routine - Unscheduled   

      bid  timestamp  year                         name  \
0  100010 2019-04-03  2019        ILLY CAFFE SF_PIER 39   
1  100017 2019-08-16  2019  AMICI'S EAST COAST PIZZERIA   
2  100041 2019-05-20  2019               UNCLE LEE CAFE   
3  100055 2019-04-25  2019                Twirl and Dip   
4  100055 2019-09-12  2019                Twirl and Dip   

                         address  
0               PIER 39  K-106-B  
1                    475 06th St  
2                 3608 BALBOA ST  
3  335 Martin Luther King Jr. Dr  
4  335 Martin Luther King Jr
```

```python
assert("name" in ins_named and "address" in ins_named) 
```

```python
grader.check("q4a")
```

Output:
```
q4a results: All test cases passed!
```

<br/><br/><br/>

--- 

## Question 4b

Let's look at the 20 businesses with the lowest **median** score.  Order your results by the median score followed by the business name to break ties. The resulting table should look like:


*Hint: You may find the `as_index` argument in the `groupby` method important*. [The documentation is linked here!](https://pandas.pydata.org/pandas-docs/stable/reference/api/pandas.DataFrame.groupby.html)

<table border="1" class="dataframe">  <thead>    <tr style="text-align: right;">      <th>bid</th>      <th>name</th>      <th>median score</th>    </tr>  </thead>  <tbody>    <tr>      <th>84590</th>      <td>Chaat Corner</td>      <td>54.0</td>    </tr>    <tr>        <th>90622</th>      <td>Taqueria Lolita</td>      <td>57.0</td>    </tr>    <tr>         <th>94351</th>      <td>VBowls LLC</td>      <td>58.0</td>    </tr>    <tr>          <th>69282</th>      <td>New Jumbo Seafood Restaurant</td>      <td>60.5</td>    </tr>    <tr>         <th>1154</th>      <td>SUNFLOWER RESTAURANT</td>      <td>63.5</td>    </tr>  <tr>          <th>93150</th>      <td>Chez Beesen</td>      <td>64.0</td>    </tr>   <tr>     <th>39776</th>      <td>Duc Loi Supermarket</td>      <td>64.0</td>    </tr>  <tr>         <th>78328</th>      <td>Golden Wok</td>      <td>64.0</td>    </tr>  <tr>          <th>69397</th>      <td>Minna SF Group LLC</td>      <td>64.0</td>    </tr>     <tr>        <th>93502</th>      <td>Smoky Man</td>      <td>64.0</td>    </tr>    <tr>           <th>98995</th>      <td>Vallarta's Taco Bar</td>      <td>64.0</td>    </tr>    <tr>         <th>10877</th>      <td>CHINA FIRST INC.</td>      <td>64.5</td>    </tr>    <tr>        <th>71310</th>      <td>Golden King Vietnamese Restaurant</td>      <td>64.5</td>    </tr>     <tr>          <th>89070</th>      <td>Lafayette Coffee Shop</td>      <td>64.5</td>    </tr>
    <tr>          <th>71008</th>      <td>House of Pancakes</td>      <td>65.0</td>    </tr> <tr>         <th>2542</th>      <td>PETER D'S RESTAURANT</td>      <td>65.0</td>    </tr>            <tr>        <th>3862</th>      <td>IMPERIAL GARDEN SEAFOOD RESTAURANT</td>      <td>66.0</td>    </tr>    <tr>         <th>61427</th>      <td>Nick's Foods</td>      <td>66.0</td>    </tr>    <tr>          <th>72176</th>      <td>Wolfes Lunch</td>      <td>66.0</td>    </tr>    <tr>        <th>89141</th>      <td>Cha Cha Cha on Mission</td>      <td>66.5</td>    </tr>  </tbody></table>


```python
twenty_lowest_scoring = ins_named.groupby(['bid', 'name'], as_index=False)['score'].median()
twenty_lowest_scoring = twenty_lowest_scoring.rename(columns={'score': 'median score'})
twenty_lowest_scoring = twenty_lowest_scoring.sort_values(['median score', 'name']).head(20)[['name', 'median score']]
twenty_lowest_scoring = twenty_lowest_scoring.set_index(ins_named.groupby(['bid', 'name'], as_index=False)['score'].median().sort_values(['score', 'name']).head(20)['bid'])

twenty_lowest_scoring
```

Output:
```
name  median score
bid                                                    
84590                        Chaat Corner          54.0
90622                     Taqueria Lolita          57.0
94351                          VBowls LLC          58.0
69282        New Jumbo Seafood Restaurant          60.5
1154                 SUNFLOWER RESTAURANT          63.5
93150                         Chez Beesen          64.0
39776                 Duc Loi Supermarket          64.0
78328                          Golden Wok          64.0
69397                  Minna SF Group LLC          64.0
93502                           Smoky Man          64.0
98995                 Vallarta's Taco Bar          64.0
10877                    CHINA FIRST INC.          64.5
71310   Golden King Vietnamese Restaurant          64.5
89070               Lafayette Coffee Shop          64.5
71008                   House of Pancakes          65.0
2542                 PETER D'S RESTAURANT          65.0
3862   IMPERIAL GARDEN SEAFOO
```

```python
assert(set(twenty_lowest_scoring.columns) == {'median score', 'name'}) 
```

```python
grader.check("q4b")
```

Output:
```
q4b results: All test cases passed!
```

### Question 4c

Let's figure out which restaurant had the worst score ever (single lowest score). 

In the cell below, assign `worst_restaurant` to the name of the restaurant with the **lowest inspection score ever**. Feel free to head to yelp.com and look up the reviews page for this restaurant.

```python
worst_restaurant = ins_named.loc[ins_named['score'].idxmin(), 'name']
worst_restaurant
```

Output:
```
'Lollipot'
```

```python
assert(type(worst_restaurant) == str) 
assert(len(worst_restaurant) > 0)
```

```python
grader.check("q4c")
```

Output:
```
q4c results: All test cases passed!
```

## Congratulations! You have finished Assignment 2! ##

Tests for q0 (points: 0)


Tests for q1a (points: 1)
```
>>> assert is_bid_unique
```


Tests for q1b (points: 2)
```
>>> assert len(top_names) == 5
```
```
>>> assert top_names[0] == "Peet's Coffee & Tea"
```


Tests for q1c (points: 1)
```
>>> assert q1c in set(['A', 'B', 'C'])
```


Tests for q2a (points: 2)
```
>>> assert type(zip_counts) == pd.Series
```
```
>>> assert zip_counts['94103'] == 562
```


Tests for q2b (points: 1)
```
>>> assert q2b in set(['A', 'B', 'C'])
```


Tests for q2ci (points: 1)
```
>>> assert (valid_zips.dtype == object) | (valid_zips.dtype == pd.StringDtype())
```
```
>>> assert type(valid_zips) == pd.Series
```


Tests for q2cii (points: 1)
```
>>> assert type(invalid_zip_bus) == pd.DataFrame
```


Tests for q2d (points: 2)
```
>>> assert type(missing_zip_address_count) == pd.Series
```
```
>>> assert missing_zip_address_count['3914 Judah St'] == 1
```


Tests for q2e (points: 2)
```
>>> assert 'postal5' in bus.columns
```
```
>>> assert (bus['postal5'].str.len() != 5).sum() == 221
```


Tests for q3a (points: 1)
```
>>> assert type(is_ins_iid_a_primary_key) == bool or type(is_ins_iid_a_primary_key) == np.bool
```


Tests for q3bi (points: 2)
```
>>> assert 'bid' in ins.columns
```
```
>>> assert ins['bid'].dtype == int
```


Tests for q3bii (points: 2)
```
>>> assert invalid_bid_count == 0
```


Tests for q3ci (points: 1)
```
>>> assert type(ins_date_type) == type
```
```
>>> assert ins_date_type in [type('str'), type('int')]
```


Tests for q3cii (points: 1)
```
>>> assert type(ins['timestamp'][1]) == pd.Timestamp
```


Tests for q3ciii (points: 1)
```
>>> assert type(earliest_date) == pd.Timestamp
```
```
>>> assert type(latest_date) == pd.Timestamp
```


Tests for q3civ (points: 1)
```
>>> assert 'year' in ins.columns
```


Tests for q4a (points: 2)
```
>>> assert 'name' in ins_named and 'address' in ins_named
```


Tests for q4b (points: 3)
```
>>> assert set(twenty_lowest_scoring.columns) == {'median score', 'name'}
```


Tests for q4c (points: 2)
```
>>> assert type(worst_restaurant) == str
```
```
>>> assert len(worst_restaurant) > 0
```
